---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Lesson 1:  Introduction to Data Engineering

```yaml
type: VideoExercise
key: 2719c20b06
xp: 50
```

`@projector_key`
916367ae330aac6a7e38d6bdbe4ff426
